```python
import os
import random
import pandas as pd
import numpy as np
import seaborn as sns
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import (
    mean_absolute_error,
    mean_absolute_percentage_error,
    mean_squared_error,
    root_mean_squared_error,
    r2_score,
)
import torch
import torch.nn as nn
import torch.nn.functional as F
import torch.optim as optim
from torch.utils.data import TensorDataset, DataLoader
```

```python
def set_seed(seed: int = 101) -> None:
    random.seed(seed)
    np.random.seed(seed)
    torch.manual_seed(seed)
    os.environ["PYTHONHASHSEED"] = str(seed)
set_seed()
```

```python
df = pd.read_csv("healthcare.csv")
```

```python
df.head()
```

Output:
```
timestamp  admissions  avg_length_of_stay  bed_occupancy_rate  \
0      18362        76.0            2.466166            0.641481   
1      18363        74.0            2.457181            0.620210   
2      18364        68.0            2.416470            0.599630   
3      18365        64.0            2.445352            0.562957   
4      18366        62.0            2.405359            0.529342   

   staff_workload  adverse_events_rate  
0        2.009302                 12.0  
1        2.005025                 13.0  
2        2.014143                 13.0  
3        1.927639                 13.0  
4        1.883696                 13.0
```

<!-- grade_id: cell-34c037f33ff030b3, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def convert_to_datetime(column: pd.Series) -> pd.Series:
    """
    Convert the values of the given `column` from days (`D`) since the epoch to `datetime64[ns]`.
    Return `column`.
    
    Note:
        Remove the `raise NotImplementedError()` line.
    """

    ### BEGIN SOLUTION
    # df["timestamp"] = pd.to_datetime(df["timestamp"], unit="D") PAZI TREBA COLUMN DVA PATI,ako ne dava error toa ns zanci treba kolona da stavis
    column= pd.to_datetime(column, unit="D")

    ### END SOLUTION

    return column
```

<!-- grade_id: cell-9aaf8653c4c7b369, points: 1, grade: true, solution: false -->
```python
### TEST
def test_convert_to_datetime():
    column = pd.Series([20000, 21000, 22000, 23000])
    column = convert_to_datetime(column)
    assert column.dtype == "<M8[ns]", "Make sure that the `dtype` is `datetime64[ns]`."
    assert column[0] == pd.Timestamp("2024-10-04"), "Make sure that the `unit` is correct."
    print("✔ Test passed!")
test_convert_to_datetime()
### BEGIN HIDDEN TESTS
def hidden_test_convert_to_datetime():
    column = pd.Series([20100, 21100, 22100, 23100])
    column = convert_to_datetime(column)
    assert column.dtype == "<M8[ns]", "Make sure that the `dtype` is `datetime64[ns]`."
    assert column[0] == pd.Timestamp("2025-01-12"), "Make sure that the `unit` is correct."
    print("✔ Test passed!")
hidden_test_convert_to_datetime()
### END HIDDEN TESTS
```

Output:
```
✔ Test passed!
✔ Test passed!
```

```python
df["timestamp"] = convert_to_datetime(df["timestamp"])
```

```python
df.head()
```

Output:
```
timestamp  admissions  avg_length_of_stay  bed_occupancy_rate  \
0 2020-04-10        76.0            2.466166            0.641481   
1 2020-04-11        74.0            2.457181            0.620210   
2 2020-04-12        68.0            2.416470            0.599630   
3 2020-04-13        64.0            2.445352            0.562957   
4 2020-04-14        62.0            2.405359            0.529342   

   staff_workload  adverse_events_rate  
0        2.009302                 12.0  
1        2.005025                 13.0  
2        2.014143                 13.0  
3        1.927639                 13.0  
4        1.883696                 13.0
```

```python
df.set_index(keys=["timestamp"], drop=True, inplace=True)
```

```python
df.head()
```

Output:
```
admissions  avg_length_of_stay  bed_occupancy_rate  \
timestamp                                                        
2020-04-10        76.0            2.466166            0.641481   
2020-04-11        74.0            2.457181            0.620210   
2020-04-12        68.0            2.416470            0.599630   
2020-04-13        64.0            2.445352            0.562957   
2020-04-14        62.0            2.405359            0.529342   

            staff_workload  adverse_events_rate  
timestamp                                        
2020-04-10        2.009302                 12.0  
2020-04-11        2.005025                 13.0  
2020-04-12        2.014143                 13.0  
2020-04-13        1.927639                 13.0  
2020-04-14        1.883696                 13.0
```

```python
features = ["avg_length_of_stay", "bed_occupancy_rate", "staff_workload", "adverse_events_rate"]
target = "admissions"
```

```python
df.isnull().sum()
```

Output:
```
admissions              0
avg_length_of_stay     33
bed_occupancy_rate     34
staff_workload         33
adverse_events_rate    37
dtype: int64
```

<!-- grade_id: cell-a559087ef66c207a, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def impute_missing_values(column: pd.Series) -> pd.Series:
    """
    Impute the missing values in the given `column`
    using forward fill and backward fill (in the specified order).
    Return `column`.
    
    Note:
        Remove the `raise NotImplementedError()` line.
    """

    ### BEGIN SOLUTION
    #df.interpolate(inplace=True) //linaer e voa a nas bfil i ffil i COLUUMN VIKA

    #nauci na pamat
    column=column.ffill().bfill()
    ### END SOLUTION

    return column
```

<!-- grade_id: cell-f9c015a2c70d3b35, points: 1, grade: true, solution: false -->
```python
### TEST
def test_impute_missing_values():
    column = pd.Series([np.nan, 2, np.nan, 4])
    column = impute_missing_values(column)
    assert column[0] == 2 and column[2] == 2, "Make sure you are using forward and backward fill for data imputation."
    print("✔ Test passed!")
test_impute_missing_values()
### BEGIN HIDDEN TESTS
def hidden_test_impute_missing_values():
    column = pd.Series([4, np.nan, 6, np.nan])
    column = impute_missing_values(column)
    assert column[1] == 4 and column[3] == 6, "Make sure you are using forward and backward fill for data imputation."
    print("✔ Test passed!")
hidden_test_impute_missing_values()
### END HIDDEN TESTS
```

Output:
```
✔ Test passed!
✔ Test passed!
```

```python
for feature in features:
    df[feature] = impute_missing_values(df[feature])
```

```python
df.isnull().sum()
```

Output:
```
admissions             0
avg_length_of_stay     0
bed_occupancy_rate     0
staff_workload         0
adverse_events_rate    0
dtype: int64
```

<!-- grade_id: cell-82a334a09f038bb8, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def calculate_simple_moving_average(column: pd.Series, column_name: str, window_size: int) -> pd.Series:
    """
    Calculate a Simple Moving Average (SMA) for the given `column` using a window size of `window_size`.
    SMA is the average of the last `window_size` points at each point in the time series.
    Return the SMA as `pd.Series` with name `{column_name}_sma_{window_size}`.

    Example:
        If `column = [10, 20, 30, 40, 50]` and `window_size = 3`, the SMA would be calculated as:
        [NaN, NaN, (10 + 20 + 30)/3, (20 + 30 + 40)/3, (30 + 40 + 50)/3] = [NaN, NaN, 20, 30, 40].
    
    Note:
        Remove the `raise NotImplementedError()` line.
    """

    ### BEGIN SOLUTION
    raise NotImplementedError()
   
    ### END SOLUTION

    return series
```

<!-- grade_id: cell-cb67cde8ed2a39ba, points: 5, grade: true, solution: false -->
```python
### TEST
def test_calculate_simple_moving_average():
    series = pd.Series([10, 20, 30, 40, 50])
    sma = calculate_simple_moving_average(series, "A", 3)
    assert pd.isna(sma[0]) and pd.isna(sma[1]) and sma[2] == 20. and sma[3] == 30. and sma[4] == 40., "The calculation is not correct."
    assert sma.name == "A_sma_3", "The name of the series is not correct."
    print("✔ Test passed!")
test_calculate_simple_moving_average()
### BEGIN HIDDEN TESTS
def hidden_test_calculate_simple_moving_average():
    series = pd.Series([60, 70, 80, 90, 100])
    sma = calculate_simple_moving_average(series, "B", 3)
    assert pd.isna(sma[0]) and pd.isna(sma[1]) and sma[2] == 70. and sma[3] == 80. and sma[4] == 90., "The calculation is not correct."
    assert sma.name == "B_sma_3", "The name of the series is not correct."
    print("✔ Test passed!")
hidden_test_calculate_simple_moving_average()
### END HIDDEN TESTS
```

Output:
```
Cell In[173], line 4, in test_calculate_simple_moving_average()
      2 def test_calculate_simple_moving_average():
      3     series = pd.Series([10, 20, 30, 40, 50])
----> 4     sma = calculate_simple_moving_average(series, "A", 3)
      5     assert pd.isna(sma[0]) and pd.isna(sma[1]) and sma[2] == 20. and sma[3] == 30. and sma[4] == 40., "The calculation is not correct."
      6     assert sma.name == "A_sma_3", "The name of the series is not correct."

Cell In[172], line 17, in calculate_simple_moving_average(column, column_name, window_size)
      3 """
      4 Calculate a Simple Moving Average (SMA) for the given `column` using a window size of `window_size`.
      5 SMA is the average of the last `window_size` points at each point in the time series.
   (...)     13     Remove the `raise NotImplementedError()` line.
     14 """
     16 ### BEGIN SOLUTION
---> 17 raise NotImplementedError()
     19 ### END SOLUTION
     21 return series

NotImplementedError:
```

```python
for column_name in features + [target]:
    try:
        series = calculate_simple_moving_average(df[column_name], column_name, 3)
        df[series.name] = series.values
        features.append(series.name)
    except:
        pass
df.dropna(axis=0, inplace=True)
```

```python
df.head()
```

Output:
```
admissions  avg_length_of_stay  bed_occupancy_rate  \
timestamp                                                        
2020-04-10        76.0            2.466166            0.641481   
2020-04-11        74.0            2.457181            0.620210   
2020-04-12        68.0            2.416470            0.599630   
2020-04-13        64.0            2.445352            0.562957   
2020-04-14        62.0            2.405359            0.529342   

            staff_workload  adverse_events_rate  
timestamp                                        
2020-04-10        2.009302                 12.0  
2020-04-11        2.005025                 13.0  
2020-04-12        2.014143                 13.0  
2020-04-13        1.927639                 13.0  
2020-04-14        1.883696                 13.0
```

<!-- grade_id: cell-229309649d00c23d, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def create_lags(df: pd.DataFrame, features: list[str], target: str, num_lags: int) -> pd.DataFrame:
    """
    Generate lagged columns for each column (`features` and `target`), up to `num_lags`,
    within the given `df`.
    
    For each lag `k`, create new columns containing values
    shifted by `k` time steps into the past. Lagged columns are ordered
    chronologically from the most distant to the most recent lag, and the original   ***********    num,0,-1 ZNACI
    column order is preserved within each lag.

    The name of the lagged columns follows the pattern `{column}_{lag}`.

    After creating the lagged columns, drop the original feature columns,
    keeping the original target column and the lagged columns.
    
    Finally, remove the rows with missing values.
    
    Note:
        Remove the `raise NotImplementedError()` line.

    """

    ### BEGIN SOLUTION
    lags = range(num_lags,0,-1)
    for lag in lags:
        for column in features + [target]:
            df[f"{column}_{lag}"] = df[column].shift(lag)
            #bara DROP ama go ima u aud reseni
    df.dropna(axis=0, inplace=True)
    df.drop(columns=features, inplace=True)
   
    ### END SOLUTION

    return df
```

<!-- grade_id: cell-85d62ee7ddb4f4f7, points: 5, grade: true, solution: false -->
```python
### TEST
def test_create_lags():
    df = pd.DataFrame(data=[[1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 23]], columns=["A", "B", "C"])
    df = create_lags(df, ["A", "B"], "C", 2)
    assert list(df.columns) == ["C", "A_2", "B_2", "C_2", "A_1", "B_1", "C_1"] or list(df.columns[1:]) == ["A_2", "B_2", "C_2", "A_1", "B_1", "C_1"], "The order of the columns is not correct."
    assert df["A_2"][2] == 1 and df["A_2"][3] == 4, "The shifting is not correct."
    assert df["A_1"][2] == 4 and df["A_1"][3] == 7, "The shifting is not correct."
    assert df["B_2"][2] == 2 and df["B_2"][3] == 5, "The shifting is not correct."
    assert df["B_1"][2] == 5 and df["B_1"][3] == 8, "The shifting is not correct."
    assert df["C_2"][2] == 3 and df["C_2"][3] == 6, "The shifting is not correct."
    assert df["C_1"][2] == 6 and df["C_1"][3] == 9, "The shifting is not correct."
    assert all([feature not in list(df.columns) for feature in ["A", "B"]]), "The features columns are not dropped."
    assert df.isnull().sum().sum().item() == 0, "The rows with missing values are not removed."
    print("✔ Test passed!")
test_create_lags()
```

Output:
```
✔ Test passed!
```

```python
num_lags = 5
```

```python
df = create_lags(df, features, target, num_lags)
```

```python
X, y = df.drop(columns=target), df[[target]]
```

<!-- grade_id: cell-e8030bdd6d8c432c, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def generate_train_val_test_split(X: pd.DataFrame, y: pd.DataFrame) -> tuple[pd.DataFrame, pd.DataFrame, pd.DataFrame, pd.DataFrame, pd.DataFrame, pd.DataFrame]:
    """
    Generate `train_X`, `val_X`, `test_X`, `train_y`, `val_y`, and `test_y` from the given `df`.
    The ratio of train, validation, and test is 80:10:10.
    First split `df` in train and test, then further split the test into test and validation.
    Set `random_state=42` to ensure reproducibility.
    
    Note:
        Remove the `raise NotImplementedError()` line.


        train_X, test_X, train_y, test_y = train_test_split(X, y, test_size=0.2, shuffle=False, random_state=42)
    test_X, val_X, test_y, val_y = train_test_split(test_X, test_y, test_size=0.5, shuffle=False, random_state=42)
    """

    ### BEGIN SOLUTION   AssertionError: The shape of the train dataset is not correct.

    train_X, test_X, train_y, test_y = train_test_split(X, y, test_size=0.2, shuffle=False,random_state=42)
    # train_X, val_X, train_y, val_y = train_test_split(train_X, train_y, test_size=0.5, shuffle=False,random_state=42)
    #NE TREBA PAK TRAIN TUKU TSET A TOA GO NEMA NIKADE U AUD
    test_X, val_X, test_y, val_y = train_test_split(test_X, test_y, test_size=0.5, shuffle=False,random_state=42)

    
    
    ### END SOLUTION

    return train_X, val_X, test_X, train_y, val_y, test_y
```

<!-- grade_id: cell-ebda2e7fec62c7bd, points: 1, grade: true, solution: false -->
```python
### TEST
def test_generate_train_val_test_split():
    df = pd.DataFrame([[1, 2], [3, 4], [5, 6], [7, 8], [9, 10], [11, 12], [13, 14], [15, 16], [17, 18], [19, 20]])
    X, y = df[[0]], df[[1]]
    train_X, val_X, test_X, train_y, val_y, test_y = generate_train_val_test_split(X, y)
    assert train_X.shape[0] == 8 and train_y.shape[0] == 8, "The shape of the train dataset is not correct."
    assert val_X.shape[0] == 1 and val_y.shape[0] == 1, "The shape of the validation dataset is not correct."
    assert test_X.shape[0] == 1 and test_y.shape[0] == 1, "The shape of the test dataset is not correct."
    print("✔ Test passed!")
test_generate_train_val_test_split()
```

Output:
```
✔ Test passed!
```

```python
train_X, val_X, test_X, train_y, val_y, test_y = generate_train_val_test_split(X, y)
```

<!-- grade_id: cell-337fad1a0740139f, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def reshape_features(train_X: pd.DataFrame, val_X: pd.DataFrame, test_X: pd.DataFrame, num_lags: int) -> tuple[np.ndarray, np.ndarray, np.ndarray]:
    """
    Reshape `train_X`, `val_X`, `test_X` into a shape of (num_samples, num_lags, num_features).
    Return `train_X`, `val_X`, and `test_X`.
    
    Note:
        Remove the `raise NotImplementedError()` line.
    """

    ### BEGIN SOLUTION  MESTO SAMO LAG STAVU NUM_LAG
    train_X = train_X.reshape(train_X.shape[0], num_lags, (train_X.shape[1] // num_lags))
    val_X = val_X.reshape(val_X.shape[0], num_lags, (val_X.shape[1] // num_lags))
    test_X = test_X.reshape(test_X.shape[0], num_lags, (test_X.shape[1] // num_lags))
   
    ### END SOLUTION

    return train_X, val_X, test_X
```

<!-- grade_id: cell-1500be82309636bb, points: 1, grade: true, solution: false -->
```python
### TEST
def test_reshape_features():
    train_X = np.array([[1, 2, 3, 4, 5, 6], [7, 8, 9, 10, 11, 12], [13, 14, 15, 16, 17, 18], [19, 20, 21, 22, 23, 24]])
    val_X = np.array([[1, 2, 3, 4, 5, 6], [7, 8, 9, 10, 11, 12]])
    test_X = np.array([[1, 2, 3, 4, 5, 6], [7, 8, 9, 10, 11, 12]])
    train_X, val_X, test_X = reshape_features(train_X, val_X, test_X, 2)
    assert train_X.shape[0] == 4 and train_X.shape[1] == 2 and train_X.shape[2] == 3, "The shape of the train dataset is not correct."
    assert val_X.shape[0] == 2 and val_X.shape[1] == 2 and val_X.shape[2] == 3, "The shape of the validation dataset is not correct."
    assert test_X.shape[0] == 2 and test_X.shape[1] == 2 and test_X.shape[2] == 3, "The shape of the test dataset is not correct."
    print("✔ Test passed!")
test_reshape_features()
```

Output:
```
✔ Test passed!
```

```python
train_X, val_X, test_X = reshape_features(train_X.to_numpy(), val_X.to_numpy(), test_X.to_numpy(), num_lags)
```

```python
train_X = torch.tensor(train_X, dtype=torch.float32)
train_y = torch.tensor(train_y.to_numpy(), dtype=torch.float32)

val_X = torch.tensor(val_X, dtype=torch.float32)
val_y = torch.tensor(val_y.to_numpy(), dtype=torch.float32)

test_X = torch.tensor(test_X, dtype=torch.float32)
test_y = torch.tensor(test_y.to_numpy(), dtype=torch.float32)
```

```python
train_dataset = TensorDataset(train_X, train_y)
val_dataset = TensorDataset(val_X, val_y)
test_dataset = TensorDataset(test_X, test_y)
```

<!-- grade_id: cell-e239503215d1cb54, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def create_data_loader(dataset: TensorDataset) -> DataLoader:
    """
    Create a data loader for the given dataset with a batch size of 32.
    Return `data_loader`.

    Note:
        Remove the `raise NotImplementedError()` line.
    """
    ### BEGIN SOLUTION NameError: name 'data_loader' is not defined
    # train_loader = DataLoader(train_dataset, batch_size=32, shuffle=False)
    # val_loader = DataLoader(val_dataset, batch_size=32, shuffle=False)
    # test_loader = DataLoader(test_dataset, batch_size=32, shuffle=False)
    data_loader=DataLoader(test_dataset, batch_size=32, shuffle=False)

    
    ### END SOLUTION

    return data_loader
```

<!-- grade_id: cell-8c9d75d8b4a8c75d, points: 1, grade: true, solution: false -->
```python
### TEST
def test_create_data_loader():
    X = torch.randn(100, 10)
    y = torch.randint(0, 2, (100,))
    dataset = TensorDataset(X, y)
    data_loader = create_data_loader(dataset)
    assert data_loader.batch_size == 32 and isinstance(data_loader.sampler, torch.utils.data.SequentialSampler), "The data loader is not configured properly."
    print("✔ Test passed!")
test_create_data_loader()
```

Output:
```
✔ Test passed!
```

```python
train_loader = create_data_loader(train_dataset)
val_loader = create_data_loader(val_dataset)
test_loader = create_data_loader(test_dataset)
```

<!-- grade_id: cell-de20d81064abfb31, points: 0, grade: false, solution: true -->
```python
class Regressor(nn.Module):
    """
    A class for a CNN regression model.

    IMPORTANT: The constructor MUST have a `num_features` argument.
    Implement both the `__init__` and `forward` methods.

    The model should consist of the following layers:
        - a 1D convolutional layer (`conv1`) with 32 filters, a kernel of size 4, a padding of 2, and ReLU activation function;
        - a 1D average pooling layer (`pool1`) with a kernel size of 2;
        - a 1D convolutional layer (`conv2`) with 64 filters, a kernel of size 4, a padding of 2, and ReLU activation function;
        - a 1D average pooling layer (`pool2`) with a kernel size of 2;
        - a flattening layer (`flatten`),
        - a fully connected layer (`fc`) with 1 neuron.
        
    Note:
        Remove the `raise NotImplementedError()` line.
    """
    ### BEGIN SOLUTION
    def __init__(self, num_features):
        super().__init__()
        self.conv1 = nn.Conv1d(num_features, 32, kernel_size=4, padding=2)
        self.pool1 = nn.AvgPool1d(2)

        self.conv2 = nn.Conv1d(32, 64, kernel_size=4, padding=2)
        self.pool2 = nn.AvgPool1d(2)

        self.flatten = nn.Flatten()

        self.fc = nn.Linear(128, 1)

    def forward(self, x):
        x = x.permute(0, 2, 1)

        x = self.pool1(F.relu(self.conv1(x)))
        x = self.pool2(F.relu(self.conv2(x)))

        x = x.permute(0, 2, 1)

        x = self.flatten(x)

        return self.fc(x)
    ### END SOLUTION
```

<!-- grade_id: cell-80909911403b450b, points: 2, grade: true, solution: false -->
```python
### TEST
def test_regressor_1():
    model = Regressor(num_features=2)
    assert hasattr(model, "conv1")
    assert isinstance(model.conv1, nn.Conv1d)
    assert hasattr(model.conv1, "in_channels") and model.conv1.in_channels == 2
    assert hasattr(model.conv1, "out_channels") and model.conv1.out_channels == 32
    assert hasattr(model.conv1, "kernel_size") and model.conv1.kernel_size[0] == 4
    assert hasattr(model.conv1, "padding") and model.conv1.padding[0] == 2
    print("✔ Test passed!")
test_regressor_1()
```

Output:
```
✔ Test passed!
```

<!-- grade_id: cell-fad450d8940d78f0, points: 1, grade: true, solution: false -->
```python
### TEST
def test_regressor_2():
    model = Regressor(num_features=2)
    assert hasattr(model, "pool1")
    assert isinstance(model.pool1, nn.AvgPool1d)
    assert hasattr(model.pool1, "kernel_size") and model.pool1.kernel_size[0] == 2
    print("✔ Test passed!")
test_regressor_2()
```

Output:
```
✔ Test passed!
```

<!-- grade_id: cell-b342ec2989639cc5, points: 2, grade: true, solution: false -->
```python
### TEST
def test_regressor_3():
    model = Regressor(num_features=2)
    assert hasattr(model, "conv2")
    assert isinstance(model.conv2, nn.Conv1d)
    assert hasattr(model.conv2, "in_channels") and model.conv2.in_channels == 32
    assert hasattr(model.conv2, "out_channels") and model.conv2.out_channels == 64
    assert hasattr(model.conv2, "kernel_size") and model.conv2.kernel_size[0] == 4
    assert hasattr(model.conv2, "padding") and model.conv2.padding[0] == 2
    print("✔ Test passed!")
test_regressor_3()
```

Output:
```
✔ Test passed!
```

<!-- grade_id: cell-8efe3185f27a63ad, points: 1, grade: true, solution: false -->
```python
### TEST
def test_regressor_4():
    model = Regressor(num_features=2)
    assert hasattr(model, "pool2")
    assert isinstance(model.pool2, nn.AvgPool1d)
    assert hasattr(model.pool2, "kernel_size") and model.pool2.kernel_size[0] == 2
    print("✔ Test passed!")
test_regressor_4()
```

Output:
```
✔ Test passed!
```

<!-- grade_id: cell-f589d3c397c3c3e0, points: 1, grade: true, solution: false -->
```python
### TEST
def test_regressor_5():
    model = Regressor(num_features=2)
    assert hasattr(model, "flatten")
    assert isinstance(model.flatten, nn.Flatten)
    print("✔ Test passed!")
test_regressor_5()
```

Output:
```
✔ Test passed!
```

<!-- grade_id: cell-ac0a9490b91b7774, points: 3, grade: true, solution: false -->
```python
### TEST
def test_regressor_6():
    model = Regressor(num_features=2)
    x = torch.rand((10, 7, 2))
    model(x)
    print("✔ Test passed!")
test_regressor_6()
```

Output:
```
✔ Test passed!
```

```python
num_features = len(features + [target])
model = Regressor(num_features)
```

<!-- grade_id: cell-538f26014d0de331, points: 0, grade: false, solution: true -->
```python
### AUTOGRADED TASK
def define_loss_fn_and_optimizer():
    """
    Define an appropriate loss_fn and an Adam optimizer with a learning rate of 0.003.
    Return `loss_fn` and `optimizer`.

    Note:
        Remove the `raise NotImplementedError()` line.
    """
    ### BEGIN SOLUTION
    loss_fn = nn.MSELoss()
    optimizer = optim.Adam(model.parameters(), lr=3e-3)
   
    ### END SOLUTION

    return loss_fn, optimizer
```

<!-- grade_id: cell-7e0c5902914915ec, points: 2, grade: true, solution: false -->
```python
### TEST
def test_define_loss_fn_and_optimizer():
    true_y = torch.tensor([1, 2, 3], dtype=torch.float32)
    pred_y = torch.tensor([1, 4, 3], dtype=torch.float32)
    loss_fn, optimizer = define_loss_fn_and_optimizer()
    assert round(loss_fn(true_y, pred_y).item(), 2) == 1.33
    assert optimizer.param_groups[0]["lr"] == 3e-3
    print("✔ Test passed!")
test_define_loss_fn_and_optimizer()
```

Output:
```
✔ Test passed!
```

```python
loss_fn, optimizer = define_loss_fn_and_optimizer()
```

```python
def train(model, dataloader, loss_fn, optimizer):
    model.train()

    total_loss = 0
    for features, targets in dataloader:
        features, targets = features, targets

        optimizer.zero_grad()

        outputs = model(features)
        loss = loss_fn(outputs, targets)

        loss.backward()
        optimizer.step()

        total_loss += loss.item()
    
    return total_loss / len(dataloader)
```

```python
def evaluate(model, dataloader, loss_fn):
    model.eval()

    total_loss, predictions = 0, []
    with torch.no_grad():
        for features, targets in dataloader:
            features, targets = features, targets

            outputs = model(features)
            loss = loss_fn(outputs, targets)

            predictions.extend(outputs.tolist())
            
            total_loss += loss.item()

    return total_loss / len(dataloader), predictions
```

```python
num_epochs = 16
```

```python
def train_and_evaluate(model, train_loader, val_loader, loss_fn, optimizer, num_epochs):
    train_losses, val_losses = [], []
    for epoch in range(1, num_epochs + 1):
        train_loss = train(model, train_loader, loss_fn, optimizer)
        train_losses.append(train_loss)
        val_loss, pred_y = evaluate(model, val_loader, loss_fn)
        val_losses.append(val_loss)

        val_mape = mean_absolute_percentage_error(val_y.tolist(), pred_y)

        print(f"Epoch {epoch:3d}/{num_epochs} | train_loss={train_loss:.4f} val_loss={val_loss:.4f} val_mape={val_mape:.4f}")
    
    return train_losses, val_losses
```

```python
train_losses, val_losses = train_and_evaluate(model, train_loader, val_loader, loss_fn, optimizer, num_epochs)
```

Output:
```
File ~\PycharmProjects\vnpVezbi\.venv\Lib\site-packages\sklearn\metrics\_regression.py:102, in _check_reg_targets(y_true, y_pred, multioutput, dtype)
     68 def _check_reg_targets(y_true, y_pred, multioutput, dtype="numeric"):
     69     """Check that y_true and y_pred belong to the same regression task.
     70 
     71     Parameters
   (...)    100         correct keyword.
    101     """
--> 102     check_consistent_length(y_true, y_pred)
    103     y_true = check_array(y_true, ensure_2d=False, dtype=dtype)
    104     y_pred = check_array(y_pred, ensure_2d=False, dtype=dtype)

File ~\PycharmProjects\vnpVezbi\.venv\Lib\site-packages\sklearn\utils\validation.py:457, in check_consistent_length(*arrays)
    455 uniques = np.unique(lengths)
    456 if len(uniques) > 1:
--> 457     raise ValueError(
    458         "Found input variables with inconsistent numbers of samples: %r"
    459         % [int(l) for l in lengths]
    460     )

ValueError: Found input variables with inconsis
```

```python
sns.lineplot(x=range(1, num_epochs + 1), y=train_losses, label="train_loss")
sns.lineplot(x=range(1, num_epochs + 1), y=val_losses, label="val_loss")
```

```python
test_loss, pred_y = evaluate(model, test_loader, loss_fn)
```

```python
print(f"MAE: {mean_absolute_error(test_y, pred_y)}")
print(f"MAPE: {mean_absolute_percentage_error(test_y, pred_y)}")
print(f"MSE: {mean_squared_error(test_y, pred_y)}")
print(f"RMSE: {root_mean_squared_error(test_y, pred_y)}")
print(f"R2 Score: {r2_score(test_y, pred_y)}")
```

Output:
```
MAE: 55.21469861651779
MAPE: 0.6756819511190675
MSE: 3121.5177422581037
RMSE: 55.870544495808375
R2 Score: -23.494417288943314
```

<!-- grade_id: cell-a0f1209d325c31ee, points: 3, grade: true, solution: false -->
```python
### TEST
assert r2_score(test_y, pred_y) > 0.85
print("✔ Test passed!")
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[230], line 2
      1 ### TEST
----> 2 assert r2_score(test_y, pred_y) > 0.85
      3 print("✔ Test passed!")

AssertionError:
```

```python

```

```python

```